CRB3 navigates Rab11 trafficking vesicles to promote γTuRC assembly during ciliogenesis

Author Response

The following is the authors’ response to the previous reviews

Reviewer # 1 (Public Review)

Specific comments

1. For all cell-based assays using shRNA to knock down CRB3, it would be desirable to perform rescue experiments to ensure that the observed phenotype of CRB3 depleted cells is specific and not due to off-target effects of the shRNA.

Thank you for your comments. Based on your suggestions, we performed the rescue experiments to observe any alterations in the primary cilia of CRB3-depleted MCF10A cells with overexpressed CRB3. The revised parts can be found in lines 186-188 and the new Supplementary Figure 3A-C has been added.

1. Figure 3G: it is very difficult to see that the red stained structures are primary cilia.

Yes, the staining structure of primary cilia in mammary ductal lumen are less clear than that of individual cells and in renal tubule in Figure 3G. We used recognized acetylated tubulin and γ-tubulin to stain the primary cilia, which were clearly labeled in individual cells. However, the labeled primary cilia in renal tubule were longer length and demonstrated a more pronounced structure than those in the mammary ductal lumen. In the mammary ductal lumen of the 10 mice we analyzed, the primary cilia showed shorter length and staining structure than the others shown in Figure 3G. This difference may be due to the distinct characteristics of primary cilia in different tissues.

1. Figure 5A: it is unfortunate the authors chose not to show the original dataset (Excel file) used for generating this figure; this makes it difficult to interpret the data. It is general policy of the journal to make source data accessible to the scientific community.

In accordance with the journal policy, we have provided the original dataset (Excel file) for Figure 5A, as detailed in “Figure 5–Source Data 1”.

1. The authors have a tendency to overinterpret their data, and not all claims put forth by the authors are fully supported by the data provided.

We have carefully read through the whole text and have revised the overinterpretation parts. These parts can be found in lines 48-50, lines 93-95, and lines 260-261.

Reviewer # 2 (Public Review)

Thank you for recognizing and supporting our research for this manuscript.

Reviewer # 1 (Recommendations For The Authors)

1. Abstract line 48-51: data overinterpretation. The authors cannot claim this based on the data they are presenting. Please modify the statement/temper the claims.

Thanks for your comments. We have revised this sentence in the abstract, as well as lines 48-50 for details.

1. There are several grammatical errors throughout the manuscript. In particular, the following sentences/statements are either wrong, confusing or non-sensical: lines 55-56; lines 87-90; lines 93-95; lines 385-387; lines 409-410.

Thanks for your positive comments. We have modified lines 55-56 to become new lines 54-55. These sentences in lines 87-90 and lines 93-95 are difficult to understand and logically problematic, so we have carefully revised this paragraph (new lines 85-90). Lines 385-387 have been deleted as they are non-sensical. Lines 409-410 contain misrepresentations. We have revised them in new lines 408-409.

1. Lines 257-259: this is data over-interpretation. It is not correct to state CRB3 is highly dynamic without having done any live cell imaging.

Thank you for your comments. We have revised this sentence, see revised lines 260-261 for details.

1. Figure 8E: if cells do not make cilia when CRB3 is lost (Figure 3), how is it possible to analyze SMO localization to cilia in these cells?

Thank you for your comments. We used immunofluorescence techniques, with acetylated tubulin and SMO co-staining, to analyze the localization of SMO to cilia. The results of immunofluorescent staining of primary cilium and statistical analysis in Figure 3 showed that the proportion of cells with primary cilium was significantly lower in the CRB3 knockdown group, but cells with primary cilium were still present. We used laser confocal microscopy micrographs to identify cells with primary cilium by staining acetylated tubulin, then analyzed the co-localization under the SMO channel, and finally analyzed the proportion of SMO-positive cilia. Several publications (J Cell Biol. 2020;219(6):e201904107; Science. 2008;320(5884):1777-81; Proc Natl Acad Sci U S A. 2012;109(34):13644-9.) have demonstrated that knocking down genes can affect primary cilium formation, and this method has also been used to examine the localization of SMO-related signaling pathway molecules on primary cilium.

1. Lines 366-366: based on the relative low magnification of the images in Figure 8H it is difficult to assess the subcellular localization of GLI1 and whether there is a difference between wild type and the Crb3 mutant cells. For example, it is not clear if GLI1 is localizing to the centrosome-cilium axis. Please modify the text accordingly.

Thank you for your good suggestions. As you mentioned, IHC cannot observe the subcellular localization of GLI1 on the centrosome-cilium axis. However, since GLI1 is a transcriptional effector at the terminal end of the Hh signaling pathway, we may not have made it clear that what we observed in the IHC results was the localization of GLI1 in the nucleus. Therefore, we have revised the description accordingly, as described in line 368 and lines 520-521.

1. Figure 7D, E: the zoomed-in images look pixelated.

Thank you for your positive comments. We have replaced these images in the new Figure 7D and E.

1. Figure 8B: Acetylacte-tub is misspelled.

Thank you for your comments. We have revised and standardized the acetylated tubulin stain to "Ace-tubulin" in all immunofluorescent images throughout the manuscript.

Reviewer # 2 (Recommendations For The Authors)

1.

CRB3 is present in mammals as 2 isoforms, A and B, originating from an alternative splicing. In this study, the authors never mention this fact and when using approaches to KO or KD CRB3A/B they are likely to deplete both isoforms which have been shown to have different C-terminal domains and functions (Fan et al., 2007). This is also important for the CRB3 antibodies used in the study since according to the material and methods section they are either against the extracellular domain common to both isoforms or the intracellular domain which is only similar in the domain close to transmembrane between the 2 isoforms. Since the antibodies used in each figure are not detailed it is impossible to know if the authors are detecting CRB3A or B or both. Please provide the information and correct for the actual isoform detected in the data and conclusions.

From the revised version we know now that CRB3B is used for exogenous expression. It has been shown that each isoform has a different role and localization in cells so why focus only on CRB3B for this study?

Thank you for your positive comments. First, previous literature has reported that CRB3b localizes in the primary cilium of MDCK cells. We have corrected the Introduction to specify CRB3b (line 81). Secondly, in the methodology section, we show that the CDs sequence of CRB3b was PCR-amplified from RNA extracted from MCF10A cells. We also designed primers specific to CRB3a but were unable to amplify them, indicating that CRB3b is significantly more expressed in epithelial cells than CRB3a. Finally, according to the company recommended by Genecards website for purchasing CRB3 cloning products, the only CRB3 sequence available in the CRB3 cDNA ORF Clone in Cloning Vector, Human (Cat: HG14324-G) from Sino Biological is CRB3b.

1.

The authors use GFP-CRB3A/B, it is not stated which isoform, over-expression to localize CRB3A/B in MCF10A cells (figure 4A). The levels of expression appear to be very high in the GFP panel and it is likely that the secretory pathway of the cells is clogged with GFP-CRB3A/B in transit from the ER to the plasma membrane. Thus, the colocalization with pericentrin might be due to the accumulation of ER and Golgi around the centrosome. This colocalization should be done with the endogenous CRB3A/B and with a better resolution.

The authors do not answer about the potential mislocalization of overexpressed exogenous protein.

We acknowledge the reviewer's perspective. The large amount of exogenous protein overexpression in the cell could potentially obstruct the protein secretion pathway, resulting in the accumulation of the exogenous protein at the ER and Golgi. Such accumulation could create the false impression of co-localization between CRB3b and the centrosome. To provide additional details (lines 215-217 and lines 426-433), we re-expressed the results exogenously and subsequently used staining of endogenous CRB3 and γ-tubulin in Fig. 4C to confirm the co-localization of CRB3 and the centrosome.

1.

The staining for CRB3A/B in Figure 4C (red) is striking with a very strong accumulation in an undefined intracellular structure and the authors do not provide any explanation for such a difference with the GFP-CRB3A/B just above.

The authors explain that two different photonic techniques are used (classical versus confocal) but in a cell biology manuscript confocal microscopy is now the standard technique.

Thank you for your comments. We have included a discussion on the partial concordance between CRB3's endogenous staining and exogenous expression results in the "Discussion" section, specifically in lines 420-435.

1.

In addition, the authors claim (Line 251/252) that Rab11 is necessary for the transport of CRB3A/B but they should KD Rab11 to show this.

The author's answer is that blocking endocytosis with dynasore is as good as knocking down Rab11 to show its interaction and role in CRB3A/B transport which is not the case.

Thank you for your comments. As requested by the reviewers, we have conducted experiments to knockdown Rab11 and detect CRB3 intracellular trafficking, as shown in the new Supplementary Figure 5B and added lines 258-260. These results provide additional support for our conclusions.

1.

The domain of CRB3A/B that is necessary for the interaction with Rab11 is the N-terminal part of the extracellular domain. This domain is thus inside the transport vesicles and not accessible from the cytoplasm. Given that Rab11 is a cytoplasmic protein, how the 2 proteins could interact across the membrane? The authors do not even discuss this essential point for their hypothesis. Comment on the revised version: the authors still do not understand the basic of cell biology since they claim that the extracellular domain of CRB3 can be in contact with Rab11 after endocytosis. Even after endocytosis the extracellular domain of CRB3A/B is inside the lumen of the endosome and not in contact with the cytosol where Rab11 is located. Lines 420-421 of the revised manuscript still claim this interaction between the two proteins without providing the link between the cytosol where Rab11 is and the endosome lumen where the extracellular domain of CRB3A/B is. Please correct.

Thank you for your positive comments. After carefully studying the relevant knowledge, we strongly agree with the reviewer's point of view. We have toned down our claim and removed the description regarding the binding of Rab11 endosomes to specific structural domains of intracellular CRB3 that we were unable to confirm (see lines 443-444 and lines 465-466).

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The primary cilium plays important roles in regulating cell differentiation, signal transduction, and tissue organization. Dysfunction of the primary cilium can lead to ciliopathies and cancer. The formation and organization of the primary cilium are highly associated with cell polarity proteins, such as the apical polarity protein CRB3. However, the molecular mechanisms by which CRB3 regulates ciliogenesis and the location of CRB3 remain unknown. Here, we show that CRB3, as a navigator, regulates vesicle trafficking in γ-TuRC assembly during ciliogenesis and cilium-related Hh and Wnt signaling pathways in tumorigenesis. Crb3 knockout mice display severe defects of the primary cilium in the mammary ductal lumen and renal tubule, while mammary epithelial-specific Crb3 knockout mice exhibit promotion of ductal epithelial hyperplasia and tumorigenesis. CRB3 is essential for lumen formation and ciliary assembly in the mammary epithelium. We demonstrate that CRB3 localizes to the basal body and that CRB3 trafficking is mediated by Rab11-positive endosomes. Significantly, CRB3 interacts with Rab11 to navigate GCP6/Rab11 trafficking vesicles to CEP290, resulting in intact γ-TuRC assembly. In addition, CRB3-depleted cells are unresponsive to the activation of the Hh signaling pathway, while CRB3 regulates the Wnt signaling pathway. Therefore, our studies reveal the molecular mechanisms by which CRB3 recognizes Rab11-positive endosomes to navigate apical vesicle trafficking in effective ciliogenesis, maintaining cellular homeostasis and tumorigenesis.

Revised version: The primary cilium plays important roles in regulating cell differentiation, signal transduction, and tissue organization. Dysfunction of the primary cilium can lead to ciliopathies and cancer. The formation and organization of the primary cilium are highly associated with cell polarity proteins, such as the apical polarity protein CRB3. However, the molecular mechanisms by which CRB3 regulates ciliogenesis and the location of CRB3 remain unknown. Here, we show that CRB3, as a navigator, regulates vesicle trafficking in γ-TuRC assembly during ciliogenesis and cilium-related Hh and Wnt signaling pathways in tumorigenesis. Crb3 knockout mice display severe defects of the primary cilium in the mammary ductal lumen and renal tubule, while mammary epithelial-specific Crb3 knockout mice exhibit the promotion of ductal epithelial hyperplasia and tumorigenesis. CRB3 is essential for lumen formation and ciliary assembly in the mammary epithelium. We demonstrate that CRB3 localizes to the basal body and that CRB3 trafficking is mediated by Rab11-positive endosomes. Significantly, CRB3 interacts with Rab11 to navigate GCP6/Rab11 trafficking vesicles to CEP290, resulting in intact γ-TuRC assembly. In addition, CRB3-depleted cells are unresponsive to the activation of the Hh signaling pathway, while CRB3 regulates the Wnt signaling pathway. Therefore, our studies reveal the molecular mechanisms by which CRB3 recognizes Rab11-positive endosomes to facilitate ciliogenesis, and regulates cilium-related signaling pathways in tumorigenesis.

Introduction

Earlier version: Epithelial tissues are the most common and widely distributed tissues in the body, and they perform several particular functions. The integrity of the cellular architecture composed of epithelial cells is necessary for these functions. Establishing apical-basal cell polarity within epithelial cells is a crucial biological process for epithelial homeostasis. Unfortunately, the loss of apical-basal cell polarity is one of the most significant hallmarks of advanced malignant tumors[1], and several reports have shown that the loss of some polarity proteins causes disordered cell polarity within epithelial cells, which cannot tightly control cellular growth and migration, ultimately leading to tumor formation, progression, and metastasis[2, 3]. However, the relationship between polarity proteins and tumorigenesis remains incomplete.

Revised version: Epithelial tissues are the most common and widely distributed tissues in the body, and they perform several particular functions. Epithelial tissues are made up of epithelial cells. Establishing apical-basal cell polarity within epithelial cells is a crucial biological process for epithelial homeostasis. Unfortunately, the loss of apical-basal cell polarity is one of the most significant hallmarks of advanced malignant tumors[1], and several reports have shown that the loss of some polarity proteins causes disordered cell polarity within epithelial cells, which cannot tightly control cellular growth and migration, ultimately leading to tumor formation, progression, and metastasis[2, 3]. However, the relationship between polarity proteins and tumorigenesis remains incomplete.

The Crumb complex, located apically, is vital in regulating and maintaining apical-basal cell polarity within epithelial cells. In mammals, three Crumbs orthologs are CRB1, CRB2, and CRB3, and only CRB3 is widely expressed in epithelial cells[4]. Our previous studies have shown that CRB3 is expressed at low levels in renal and breast cancers, and abnormal CRB3 expression leads to the disrupted organization of MCF10A cells in three-dimensional (3D) cultures[5–7]. Inhibition of CRB3 impairs contact inhibition and leads to migration, invasion, and tumorigenesis of cancer cells[6, 8]. Contact inhibition results in growth arrest and causes epithelial cells to enter the quiescent phase, inducing primary cilium formation.

Earlier version: The primary cilium, an antenna-like microtubule-based organelle, is a sensorial antenna that regulates cell differentiation, proliferation, polarity, and tissue organization in most types of mammalian cells[9]. Dysfunction of the primary cilium can lead to developmental and degenerative disorders called ciliopathies, such as polycystic kidney disease, nephronophthisis, retinitis pigmentosa, and Meckel syndrome[10]. Remarkably, many cancers, including melanoma and breast, renal, pancreatic, and prostate cancer, exhibit loss of the primary cilium, most likely during the early stages of tumorigenesis[11–15]. The primary cilium is a significant extension on the apical surfaces of epithelial cells and is maintained by polarized vesicular traffic[16]. Studies have shown that CRB3 localizes in the primary cilium and that depletion of CRB3 leads to primary cilium loss in Madin-Darby canine kidney (MDCK) cells[17, 18], indicating that CRB3 is necessary for the initiation of ciliogenesis. However, how CRB3 promotes ciliogenesis on the apical surfaces of epithelial cells remains unknown.

Revised version: The primary cilium, an antenna-like microtubule-based organelle, is a sensorial antenna that regulates cell differentiation, proliferation, polarity, and tissue organization in most types of mammalian cells[9]. Dysfunction of the primary cilium can lead to developmental and degenerative disorders called ciliopathies, such as polycystic kidney disease, nephronophthisis, retinitis pigmentosa, and Meckel syndrome[10]. Remarkably, many cancers, including melanoma and breast, renal, pancreatic, and prostate cancer, exhibit loss of the primary cilium, most likely during the early stages of tumorigenesis[11–15]. The primary cilium is a significant extension on the apical surfaces of epithelial cells and is maintained by polarized vesicular traffic[16]. Studies have shown that CRB3b localizes in the primary cilium and that depletion of CRB3 leads to primary cilium loss in Madin-Darby canine kidney (MDCK) cells[17, 18], indicating that CRB3 is necessary for the initiation of ciliogenesis. However, how CRB3 promotes ciliogenesis on the apical surfaces of epithelial cells remains unknown.

Earlier version: Here, we describe a novel conditional deletion of Crb3 in mice and show that Crb3 is required to assemble the primary cilium in the mammary gland, kidney and MEF cells from Crb3 knockout mice and that Crb3 deletion promotes breast cancer progression in vivo. Additionally, we found that CRB3 interacts with Rab11 and navigates GCP6/Rab11 trafficking vesicles to the basal body of the primary cilium in mammary epithelial cells. Furthermore, we identified that CRB3 regulates the ciliary Hedgehog (Hh) and Wnt signaling pathways in breast cancer. Thus, there is unexpected diversity in the polarity protein CRB3 navigating polarized Rab11-dependent traffic, leading to a novel assembly mechanism of the γ-tubulin ring complex (γTuRC) in ciliogenesis.

Revised version: Here, we describe a novel conditional deletion of Crb3 in mice. Our results demonstrate that Crb3 is required to assemble the primary cilium in the mammary gland, kidney, and MEF cells from Crb3 knockout mice. Moreover, we found that Crb3 deletion promotes breast cancer progression in vivo. Additionally, we observed that CRB3 knockdown destabilizes the assembly of the γ-tubulin ring complex (γTuRC) during ciliogenesis. Specifically, CRB3 interacts with Rab11 and navigates GCP6/Rab11 trafficking vesicles to the basal body of the primary cilium in mammary epithelial cells. We also identified that CRB3 regulates the ciliary Hedgehog (Hh) and Wnt signaling pathways in breast cancer. Thus, the polarity protein CRB3 is important for the assembly mechanism of the γTuRC in ciliogenesis and the cilium-related signaling pathways in tumorigenesis.

Results

Crb3 deletion mice exhibit smaller size and anophthalmia

Since Crb3 knockout mice die shortly after birth[19], we generated a novel transgenic mouse model with conditional deletion of Crb3 using the Cre-loxP system. The loxP sites flanked either side of exon 3 in the Crb3 gene (Supplementary Fig. 1A). In mammals, CRB3 has two isoforms, CRB3a and CRB3b, distinguished by alternative splicing within the fourth exon of the CRB3 gene, which in turn produces a protein with 23 amino acid differences at the C terminus[17]. According to the construction strategy, we are knocking out both isoforms Crb3a and Crb3b in the mice. The positive embryonic stem (ES) cell clone was confirmed for the wild-type gene (Crb3wt/wt) and recombinant allele (Crb3wt/fl) by Southern blotting (Supplementary Fig. 1B). The genotypes of the wild type, heterozygotes, and homozygotes (Crb3fl/fl) were detected with specific primers (Supplementary Fig. 1C). We first intercrossed Crb3fl/fl with CMV enhancer/chicken β-actin promoter (CAG)-Cre mice and found that Crb3fl/fl, Crb3wt/fl; CAG-Cre and Crb3fl/fl; CAG-Cre pups were born normally and well developed in size. However, most Crb3fl/fl; CAG-Cre pups died within a few days after birth. Few surviving Crb3fl/fl; CAG-Cre mice were smaller and showed anophthalmia than littermate Crb3fl/fl mice at 4 weeks old (Fig. 1A-B). The eye is an organ with perfect apical-basal cell polarity. However, Crb3 knockout mice showed ocular abnormalities. Together, these results indicate that Crb3 is necessary for eye development.

Crb3 knockout mice exhibit smaller sizes and ocular abnormalities, and mammary epithelial cell-specific Crb3 knockout leads to ductal epithelial hyperplasia and promotes tumorigenesis.

A, B. Representative whole bodies (A) and eyes (B) from littermate Crb3fl/fl, Crb3wt/fl;CAG-Cre and Crb3fl/fl;CAG-Cre mice at 4 weeks old. C. Representative mammary whole mounts from littermate Crb3fl/fl and Crb3fl/fl;MMTV-Cre mice at 8 weeks old with Carmine-alum staining. (scale bars, 200 μm) D, E. Quantification of the average number of TEBs and bifurcated TEBs in littermate Crb3fl/fl (n=10) and Crb3fl/fl;MMTV-Cre (n=10) mice at 8 weeks old. F. Representative images of mammary glands in littermate Crb3fl/fl and Crb3fl/fl;MMTV-Cre mice stained with H&E. (scale bars, 50 μm) G. Representative images of primary tumors stained with H&E in PyMT-WT and PyMT-cKO-Crb3 mice at 9 weeks old. (scale bars, left 500 μm, right 50 μm) The magnified areas of boxed sections are shown in the right panels. Bars represent the means ± SD; Unpaired Student’s t test, ***P<0.001.

Deletion of mammary epithelial-specific Crb3 causes ductal epithelial hyperplasia and tumorigenesis

Based on our earlier studies, which showed that inhibiting CRB3 impairs the organization of breast epithelium, we set out to directly explore the role of Crb3 in mammary gland development. We first developed epithelial cell-specific deletion of Crb3 in the mammary gland by intercrossing Crb3fl/fl mice with mouse mammary tumor virus long terminal repeat promoter (MMTV)-Cre mice. We confirmed the expression of Crb3 in the mammary gland using immunoblotting (both isoforms Crb3a and Crb3b), real-time quantitative PCR, and immunohistochemistry, and Crb3 expression was deficient in mammary epithelial cells with Crb3 knockout (Supplementary Fig. 1D-F).

To assess the function of Crb3 in the development of mammary epithelial cells, we analyzed whole mammary mounts from 8-week-old virgin Crb3fl/fl; MMTV-Cre and littermate Crb3fl/fl mice. Crb3fl/fl; MMTV-Cre mice displayed significantly increased numbers of terminal end buds (TEBs) and bifurcated TEBs compared to Crb3fl/fl mice (Fig. 1 C-E). Histopathology results revealed that normal ductal epithelial cells were arranged as a monolayer in Crb3fl/fl mice, while Crb3 knockout led to ductal epithelial hyperplasia with increased ductal thickness in Crb3fl/fl; MMTV-Cre mice (Fig. 1 F).

To investigate the role of Crb3 in breast cancer tumorigenesis, we observed mammary glands from 9-week-old virgin polyomavirus middle T antigen (PyMT)-cKO-Crb3 and PyMT-WT mice. The mammary glands of PyMT mice progressed to early carcinomas and lost their normal epithelial architecture[20]. The PyMT mouse model is widely used to study tumor initiation and development, with a histological progression similar to that of human breast cancer. PyMT-WT mice exhibited loss of epithelial features and developed breast cancer, while PyMT-cKO-Crb3 mice developed larger tumors with faster tumor progression and a poorly differentiated phenotype (Fig. 1G). Thus, these results suggest that Crb3 knockout leads to ductal epithelial hyperplasia and promotes tumorigenesis.

CRB3 knockdown inhibits lumen formation in acini of mammary epithelial cells

MCF10A cells are spontaneously immortalized human breast epithelial cells with the characteristics of normal breast epithelium[21]. In 3D basement membrane culture conditions, MCF-10A cells form acini similar to mammary epithelial acini. This makes them a valuable system for studying the morphogenesis and oncogenesis of the mammary epithelium[22]. To confirm the function of CRB3 in mammary epithelial acini formation, we first knocked down the expression of both CRB3a and CRB3b isoforms using lentiviral shRNA in MCF10A cells (Supplementary Fig. 2A and B). Under the 3D culture system, we found that MCF10A cells infected with a non-targeting shRNA lentivirus formed acini after 6 days (D6) and eventually formed polarized acini with a hollow lumen at D14 (Fig. 2A). However, after CRB3 knockdown, MCF10A cells formed more and larger aberrant acini without a lumen after D6 (Fig. 2A and B). Cell proliferation assays of MCF10A cells showed that CRB3 knockdown increased the proliferation rate and accelerated G1 to S phase progression (Supplementary Fig. 2C-E).

CRB3 knockdown inhibits acinar formation of mammary epithelial cells in a 3D culture system.

A. Representative effect of CRB3 on acinar formation in the 3D culture system at days 3, 6, 9 and 14. (The magnified areas of the marked arrows are shown in the lower right corner) B. Quantification of the average number, diameter and aberration of acini (At least 3-4 random micrographs were analyzed for each condition in each of three independent experiments, n=10). C. Immunofluorescence showing apoptosis during lumen formation. Caspase 3 (red), α-tubulin (green), and DNA (blue). D. Quantification of the proportion of caspase 3 positive acini (At least 3-4 random micrographs were analyzed for each condition in each of three independent experiments). E. Immunofluorescence showing the mitotic spindle orientation during lumen formation. α-tubulin (green) and DNA (blue). F. Quantification of division angle (At least 3-4 acini were randomly examined for each condition in each of three independent experiments). G. Immunohistochemical analyses of Ki67, phospho-histone H3 and cleaved caspase 3 in primary tumors from PyMT-WT (n=3) and PyMT-cKO-Crb3 (n=3) mice at 9 weeks old. (Positive cells are marked by arrows; scale bars, 25 μm) H. Quantification of the number of Ki67, phospho-histone H3 and cleaved caspase 3 positive cells from IHC staining images. Bars represent means ± SD; Unpaired Student’s t test, * P<0.5, ** P<0.01, ***P<0.001.

To investigate the reason for aberrant acini without lumens, we further studied apoptosis and mitotic spindle orientation during lumen formation in a 3D culture system. The immunofluorescence (IF) results showed that more internal cleaved-caspase 3-positive acini were observed in the control groups than in the CRB3 knockdown groups at D9 (Fig. 2C and D). Moreover, CRB3b overexpression significantly promoted apoptosis in breast cancer cells (Supplementary Fig. 2F-H). Similarly, the staining of α-tubulin in mitotic cells demonstrated more misorientation of the mitotic spindle in CRB3 knockdown groups (Fig. 2E and F).

In addition to verifying the effect of Crb3 deletion on promoting proliferation and inhibiting apoptosis in vivo, we detected these markers in primary tumor tissue from PyMT-WT and PyMT-cKO-Crb3 mice. Crb3 deletion increased the numbers of proliferative and mitotic cells, and decreased the numbers of apoptotic cells in primary tumors (Fig 2G and H). These results demonstrate that CRB3 could promote the proliferation of mammary epithelial cells, disrupt the mitotic spindle orientation, and protect the internal cells from apoptosis during acini formation, ultimately leading to irregular lumen formation and tumor progression.

CRB3 regulates primary cilium formation

Earlier version: The primary cilium is an apical antenna-like extensions in epithelial cells that is known to play an important role in controlling lumen formation[23]. Previous studies have reported that CRB3 knockdown leads to primary cilium loss in MDCK cells[17, 18]. However, the molecular mechanism that regulates ciliogenesis on apical surfaces is still unclear. To investigate the altered effect of CRB3 on ciliogenesis both in vivo and in vitro, we examined the formation of primary cilia in control MCF10A cells and CRB3 knockdown cells. We found that control MCF10A cells formed primary cilia after becoming confluent, while CRB3 knockdown led to significant ciliogenesis defects (Fig. 3A). The number of cells with primary cilium was significantly reduced compared to the control group, but the length of the primary cilium that did exist was not different (Fig. 3B). Conversely, CRB3b conditional overexpression by adding doxycycline (+dox) resulted in restoring ciliary assembly in MCF7 cells (Fig. 3C, Supplementary Fig. 3A). CRB3b overexpression increased the proportion of primary cilium formation in breast cancer cells, and the length of the restored primary cilium was increased (Fig. 3D). Importantly, the primary cilium could be directly visualized using scanning electron microscopy (SEM) (Fig. 3E). After CRB3 knockdown, the number of MCF10A cells with primary cilia was significantly decreased, while CRB3b conditional overexpression did not increase these populations (Fig. 3F).

Revised version: The primary cilium is an apical antenna-like extensions in epithelial cells that is known to play an important role in controlling lumen formation[23]. Previous studies have reported that CRB3 knockdown leads to primary cilium loss in MDCK cells[17, 18]. However, the molecular mechanism that regulates ciliogenesis on apical surfaces is still unclear. To investigate the altered effect of CRB3 on ciliogenesis both in vivo and in vitro, we examined the formation of primary cilia in control MCF10A cells and CRB3 knockdown cells. We found that control MCF10A cells formed primary cilia after becoming confluent, while CRB3 knockdown led to significant ciliogenesis defects (Fig. 3A). The number of cells with primary cilium was significantly reduced compared to the control group, but the length of the primary cilium that did exist was not different (Fig. 3B). The overexpression of CRB3b in CRB3-depleted MCF10A cells was able to fully restore the formation of primary cilium to levels comparable to the control group (Supplementary Fig. 3A-C). Conversely, CRB3b conditional overexpression by adding doxycycline (+dox) resulted in restoring ciliary assembly in MCF7 cells (Fig. 3C, Supplementary Fig. 3D). CRB3b overexpression increased the proportion of primary cilium formation in breast cancer cells, and the length of the restored primary cilium was increased (Fig. 3D). Importantly, the primary cilium could be directly visualized using scanning electron microscopy (SEM) (Fig. 3E). After CRB3 knockdown, the number of MCF10A cells with primary cilia was significantly decreased, while CRB3b conditional overexpression did not increase these populations (Fig. 3F).

CRB3 alters primary cilium formation in mammary cells, mammary ductal lumen and renal tubule from Crb3fl/fl;CAG-Cre mice.

A. C. Representative images of immunofluorescent staining of primary cilium formation with CRB3 knockdown in MCF10A cells and CRB3b conditional overexpression upon dox induction in MCF7 cells. Acetylated tubulin (red), γ-tubulin (green), and DNA (blue). (The primary cilium is marked by arrows; scale bars, 10 μm) B. D. Quantification of the proportion of cells with primary cilium formation and the length of the primary cilium (At least 3-4 random micrographs were analyzed for each condition in each of three independent experiments, n=10). E. Representative scanning electron microscope images of primary cilium formation with CRB3 knockdown and CRB3b conditional overexpression in MCF10A cells. (The primary cilium is marked by arrows; scale bars, 50 μm) F. Quantification of the proportion of cells with primary cilium formation (At least 3-4 random micrographs were analyzed for each condition in each of three independent experiments, n=10). G. H. Representative immunofluorescent staining of primary cilium formation in the mammary ductal lumen and renal tubule from Crb3fl/fl (n=10) and Crb3fl/fl;CAG-Cre (n=10) mice, respectively. Acetylated tubulin (red), γ-tubulin (green), and DNA (blue). (scale bars, 25 μm) Bars represent the means ± SD; Unpaired Student’s t test, *P<0.05, ***P<0.001.

Earlier version: To further support these findings in vivo, we used breast and kidney tissue to assess ciliogenesis in the Crb3 knockout mouse model. Renal epithelial cells can form prominent primary cilia, and the absence of primary cilia leads to ciliopathies. We used immunofluorescence to detect ciliary formation in the mammary ductal lumen and renal tubule. Compared to tissues of Crb3fl/fl mice, Crb3 knockout led to the absence of the primary cilium in the mammary ductal lumen and renal tubule from Crb3fl/fl; CAG-Cre mice (Fig. 3G, H). Specifically, the proportion of cells forming the primary cilium plummeted (Fig. 3G, H). This same phenomenon was also observed in mouse embryonic fibroblasts (MEFs) from Crb3fl/fl; CAG-Cre mice (Supplementary Fig. 3B, C). We conclude that CRB3 plays a central role in ciliogenesis in various tissues and mammary cells.

Revised version: To further support these findings in vivo, we used breast and kidney tissue to assess ciliogenesis in the Crb3 knockout mouse model. Renal epithelial cells can form prominent primary cilia, and the absence of primary cilia leads to ciliopathies. We used immunofluorescence to detect ciliary formation in the mammary ductal lumen and renal tubule. Compared to tissues of Crb3fl/fl mice, Crb3 knockout led to the absence of the primary cilium in the mammary ductal lumen and renal tubule from Crb3fl/fl; CAG-Cre mice (Fig. 3G, H). Specifically, the proportion of cells forming the primary cilium plummeted (Fig. 3G, H). This same phenomenon was also observed in mouse embryonic fibroblasts (MEFs) from Crb3fl/fl; CAG-Cre mice (Supplementary Fig. 3E, F). We conclude that CRB3 plays a central role in ciliogenesis in various tissues and mammary cells.

CRB3 localizes to the basal body of the primary cilium

Earlier version: CRB3, a polarity protein, is typically localized on the tight junctions of the apical epithelium membrane[24]. Additionally, we noted that CRB3b is found to be closely colocalized with centrosomes at prophase, metaphase, and anaphase[17]. To investigate the relationship between CRB3 and ciliary formation, we first examined CRB3 localization in MCF10A cells expressing CRB3-GFP. In MCF10A cells, exogenous CRB3b was mainly localized at cell tight junctions and could colocalize with pericentrin (a centrosome marker) (Fig. 4A). And approximately 40% of cells showed this co-localization of exogenous CRB3b and pericentrin (Fig. 4B). We verified this phenomenon by observing the localization of endogenous CRB3 (both isoforms CRB3a and CRB3b) with another centrosome marker. Similarly, we detected that CRB3 accumulated on one side of the cytoplasm and had a CRB3 foci located at the basal body represented by γ-tubulin in MCF10A cells (Fig. 4C, D). Importantly, CRB3 knockdown disturbed the accumulation of γ-tubulin at the basal body in confluent quiescence cells (Fig. 4C, E). To demonstrate the relationship of this colocalization to the primary cilium, we employed acetylated tubulin to mark the primary cilium. Double immunostaining showed that this CRB3 foci was the basal body of the primary cilium, and CRB3 knockdown inhibited ciliary assembly in MCF10A cells (Fig. 4F, G). These results suggest that CRB3 is located at the basal body of the primary cilium and that CRB3 knockdown disturbs the accumulation of γ-tubulin in quiescent cells.

Revised version: CRB3, a polarity protein, is typically localized on the tight junctions of the apical epithelium membrane[24]. Additionally, we noted that CRB3b is found to be closely colocalized with centrosomes at prophase, metaphase, and anaphase[17]. To investigate the relationship between CRB3 and ciliary formation, we first examined CRB3 localization in MCF10A cells expressing CRB3-GFP. In MCF10A cells, exogenous CRB3b was mainly localized at cell tight junctions and could colocalize with pericentrin (a centrosome marker) (Fig. 4A). And approximately 40% of cells showed this co-localization of exogenous CRB3b and pericentrin (Fig. 4B). To avoid the possibility that a large amount of exogenous overexpression leads to its accumulation at the ER and Golgi, which may affect its co-localization with the centrosome. We also verified this phenomenon by observing the localization of endogenous CRB3 (both isoforms CRB3a and CRB3b) with another centrosome marker. Similarly, we detected that CRB3 accumulated on one side of the cytoplasm and had a CRB3 foci located at the basal body represented by γ-tubulin in MCF10A cells (Fig. 4C, D). Importantly, CRB3 knockdown disturbed the accumulation of γ-tubulin at the basal body in confluent quiescence cells (Fig. 4C, E). To demonstrate the relationship of this colocalization to the primary cilium, we employed acetylated tubulin to mark the primary cilium. Double immunostaining showed that this CRB3 foci was the basal body of the primary cilium, and CRB3 knockdown inhibited ciliary assembly in MCF10A cells (Fig. 4F, G). These results suggest that CRB3 is located at the basal body of the primary cilium and that CRB3 knockdown disturbs the accumulation of γ-tubulin in quiescent cells.

CRB3 localizes to the basal body of the primary cilium.

A. Immunofluorescence showing the colocalization of exogenous CRB3b with centrosomes in MCF10A cells. Pericentrin, a marker of centrosome (red), CRB3-GFP (green), and DNA (blue). (Colocalization is marked by arrows; scale bars, 10 μm) B. Quantification of the proportion of cells with pericentrin and exogenous CRB3b colocalization (At least 3-4 random micrographs were analyzed for each condition in each of three independent experiments, n=10). C. Another colocalization of endogenous CRB3 with the basal body in MCF10A cells. γ-tubulin is a marker of the centrosome and basal body of the primary cilium (green), CRB3 (red), and DNA (blue). (Colocalization is marked by arrows; scale bars, 10 μm) D. Corresponding fluorescence intensity profile across a section of the array, as indicated by the dashed white line in (C). E. Quantification of the number of γ-tubulin foci in MCF10A cells from IF staining images (At least 3-4 random micrographs were analyzed for each condition in each of three independent experiments, n=10). F. Double immunostaining displaying the colocalization of CRB3 with the primary cilium in MCF10A cells. Acetylated tubulin (green), CRB3 (red), and DNA (blue). (Colocalization is marked by arrows; scale bars, 10 μm) G. Fluorescence 3D reconstruction of CRB3 and primary cilium colocalization. Acetylated tubulin (red), CRB3 (green), and DNA (blue). Bars represent the means ± SD; Unpaired Student’s t test, ***P<0.001.

CRB3 trafficking is mediated by Rab11-positive endosomes

To better investigate the molecular mechanisms of CRB3 in regulating primary cilium formation, we examined some important genes related to ciliogenesis. However, CRB3 knockdown did not affect the mRNA expression of these genes in MCF10A cells (Supplementary Fig. 4A). Then, we performed coimmunoprecipitation (co-IP) tandem mass spectrometry using tagged exogenous CRB3b as bait to identify its interacting proteins (Supplementary Fig. 4B). Pathway aggregation analysis of these proteins revealed that these interacting proteins of CRB3b were significantly involved in Golgi vesicle transport and vesicle organization (Fig. 5A). We focused on these pathways because smaller distal appendage vesicle (DAV) formation is critical for ciliogenesis initiation, and it requires the Rab GTPase Rab11-Rab8 cascade to function[25]. Interestingly, some Rab small GTPase family members identified as CRB3b-binding proteins were aggregated in these pathways, such as Rab10, Rab11A, Rab11B, Rab14, Rab1A, Rab1B, Rab21, Rab2A, Rab32, Rab35, Rab38, Rab5B, Rab5C, and Rab6A. Several centriolar proteins were also identified, including tubulin gamma-1, tubulin gamma-2, CENP-E, CEP290, CEP192, CEP295, and CEP162 (Fig. 5B).

CRB3 trafficking is mediated by Rab11-positive endosomes, and CRB3 knockdown destabilizes γTuRC assembly during ciliogenesis.

A. Pathway aggregation analysis of CRB3b binding proteins identified by mass spectrometry in MCF10A cells. B. Table of some Rab small GTPase family members and centriolar proteins identified as CRB3b binding proteins. C. Immunofluorescence showing the colocalization of CRB3 with EEA1-, CD63-, and Rab11-positive endosomes in MCF10A cells. EEA1, CD63, Rab11 (green), CRB3 (red), and DNA (blue). (scale bars, 10 μm) D. Quantification of the proportion of cells with these markers and CRB3 colocalization (n=10). E. Western blotting showing the levels of CRB3 in MCF10A cells treated with dynasore at different cell densities. F. The structure diagram of γTuRC (γ-tubulin ring complex). G. Immunoblot analysis of the effect of CRB3 on γTuRC molecules in MCF10A cells. H. Coimmunoprecipitation showing the interacting proteins with GCP6 in MCF10A cells with CRB3 knockdown. I. Representative images of immunofluorescent staining of GCP3 and GCP6 colocalization in MCF10A cells with the corresponding fluorescence intensity profile. GCP3 (red), GCP6 (green), and DNA (blue). (At least 3-4 random micrographs were analyzed for each condition in each of three independent experiments, scale bars, 10 μm) J. The comparison of cytoplasmic extracts from MCF10A cells and cells with CRB3 knockdown after fractionation in sucrose gradients. The γTuSC sedimentation was mainly in fractions 3, and γTuRC sedimentation was mainly in fractions 6. Bars represent the means ± SD; Unpaired Student’s t test, ***P<0.001.

Earlier version: Rab11-positive vesicles bind to centrosomal Rabin8, leading to the formation of ciliary membrane[26]. This finding suggested that the trafficking of Rab11-positive vesicles and polarized vesicles is essential for early ciliary assembly. Next, we examined the polarized vesicle traffic of CRB3 in MCF10A cells. Our studies found that CRB3 could colocalize with EEA1-positive early endosomes, CD63-positive late endosomes and Rab11-positive recycling endosomes (Fig. 5C). In addition, after treatment with dynasore, an endocytosis inhibitor, CRB3 significantly accumulated at the cell membrane after 2 hours, and the colocalization of CRB3 with EEA1-, CD63-, and Rab11-positive endosomes was significantly decreased (Fig. 5C, D). We have previously reported that CRB3 regulates cell contact inhibition, which is significantly downregulated in confluent MCF10A cells[6]. Moreover, dynasore could rescue the CRB3 downregulation in confluent MCF10A cells (Fig. 5E). The knockdown of CRB3 did not affect Rab11 expression (Supplementary Fig. 4C). These results suggest that CRB3 is highly dynamic among the endosome system, and in particular, Rab11-positive endosomes are involved in the intracellular trafficking of CRB3.

Revised version: Rab11-positive vesicles bind to centrosomal Rabin8, leading to ciliary membrane formation[26]. This finding suggested that the trafficking of Rab11-positive vesicles and polarized vesicles is essential for early ciliary assembly. Next, we examined the polarized vesicle traffic of CRB3 in MCF10A cells. Our studies found that CRB3 could colocalize with EEA1-positive early endosomes, CD63-positive late endosomes, and Rab11-positive recycling endosomes (Fig. 5C). In addition, after treatment with dynasore, an endocytosis inhibitor, CRB3 significantly accumulated at the cell membrane after 2 hours, and the colocalization of CRB3 with EEA1-, CD63-, and Rab11-positive endosomes were significantly decreased (Fig. 5C, D). We have previously reported that CRB3 regulates cell contact inhibition, which is significantly downregulated in confluent MCF10A cells[6]. Moreover, dynasore could rescue the CRB3 downregulation in confluent MCF10A cells (Fig. 5E). The knockdown of CRB3 did not affect Rab11 expression (Supplementary Fig. 4C). Rab11 knockdown resulted in a predominant localization of CRB3 to the cell membrane and a reduction in intracellular foci (Supplementary Fig. 5A, B). These results suggest that CRB3 colocalizes with EEA1, CD63, and Rab11 on endosomes, and in particular, Rab11-positive endosomes are involved in the intracellular trafficking of CRB3.

CRB3 knockdown destabilizes γTuRC assembly during ciliogenesis

Earlier version: According to our findings, we assumed that CRB3 might be involved in polarized vesicle trafficking. The γ-Tubulin ring complex (γTuRC) is a central regulator of microtubule nucleation and a nucleation site of α-Tubulin and β-Tubulin at the microtubule[organizing center in mitotic cells[27]. Additionally, the primary cilium, a microtubule-based structure attached to the basal body, is transformed from the mother centriole. As we found that CRB3 knockdown disturbed the accumulation of γ-tubulin prompted us to analyze γTuRC assembly in ciliogenesis. The γTuRC is composed of multiple copies of the γ-tubulin small complex (γTuSC), GCP2 and GCP3, plus GCP4, GCP5, and GCP6[27, 28] (Fig. 5F). Our findings indicate that CRB3 downregulation did not affect the expression of γTuRC subunits (Fig. 5G). However, we performed co-IPs directed against GCP6 in CRB3-depleted cells and assessed the relative levels of both γTuSC-specific proteins compared with the control shRNA-treated cells. Interestingly, CRB3 knockdown significantly reduced the interaction of GCP6 with GCP3 (Fig. 5H). To verify this phenomenon, we observed the localization of GCP3 with GCP6 foci in MCF10A cells. Correspondingly, CRB3 downregulation significantly disturbed the colocalization of GCP3 with GCP6 foci in quiescence cells (Fig. 5I). To investigate whether CRB3 facilitate γTuRC assembly, we fractionated MCF10A cytoplasmic proteins on sucrose gradients according to published methods[29]. We noticed that γTuSC were sedimented mainly in fractions 3, while γTuRC sedimented in fractions 6 (Fig. 5J). However, CRB3 downregulation caused most of GCP6 and GCP5 to be precipitated independently of γTuRC in the low-density fractions (Fig. 5J). Together, these results demonstrate that the depletion of CRB3 disrupts molecular interactions between the γTuRC proteins, resulting in a failure of microtubule formation in the primary cilium. However, this depletion may not affect the expression of these molecules during γTuRC assembly. Based on these findings, we hypothesize that CRB3 could facilitate the polarized vesicle trafficking of some γTuRC-specific proteins through Rab11-positive endosomes.

Revised version: According to our findings, we assumed that CRB3 might be involved in polarized vesicle trafficking. The γ-Tubulin ring complex (γTuRC) is a central regulator of microtubule nucleation and a nucleation site of α-tubulin and β-tubulin at the microtubule[organizing center in mitotic cells[27]. Additionally, the primary cilium, a microtubule-based structure attached to the basal body, is transformed from the mother centriole. As we found that CRB3 knockdown disturbed the accumulation of γ-tubulin prompted us to analyze γTuRC assembly in ciliogenesis. The γTuRC is composed of multiple copies of the γ[tubulin small complex (γTuSC), GCP2 and GCP3, plus GCP4, GCP5, and GCP6[27, 28] (Fig. 5F). Our findings indicate that CRB3 downregulation did not affect the expression of γTuRC subunits (Fig. 5G). However, we performed co-IPs directed against GCP6 in CRB3-depleted cells and assessed the relative levels of both γTuSC-specific proteins compared with the control shRNA-treated cells. Interestingly, CRB3 knockdown significantly reduced the interaction of GCP6 with GCP3 (Fig. 5H). To verify this phenomenon, we observed the localization of GCP3 with GCP6 foci in MCF10A cells. Correspondingly, CRB3 downregulation significantly disturbed the colocalization of GCP3 with GCP6 foci in quiescence cells (Fig. 5I). To investigate whether CRB3 facilitates γTuRC assembly, we fractionated MCF10A cytoplasmic proteins on sucrose gradients according to published methods[29]. We noticed that γTuSC were sedimented mainly in fractions 3, while γTuRC sedimented in fractions 6 (Fig. 5J). However, CRB3 downregulation caused most of GCP6 and GCP5 to be precipitated independently of γTuRC in the low-density fractions (Fig. 5J). Together, these results demonstrate that the depletion of CRB3 disrupts molecular interactions between the γTuRC proteins, resulting in a failure of microtubule formation in the primary cilium. However, this depletion may not affect the expression of these molecules during γTuRC assembly. Based on these findings, we hypothesize that CRB3 could facilitate the polarized vesicle trafficking of some γTuRC-specific proteins through Rab11-positive endosomes.

CRB3 interacts with Rab11

Earlier version: To fully prove this hypothesis, we first examined the interaction of CRB3 with Rab11-positive endosomes and γTuRC-specific proteins in MCF10A cells. Using CRB3 as bait, we identified that CRB3 could interact with Rab11, but not GCP6 and GCP3 (Fig. 6A). In contrast, Rab11 could interact with CRB3 and GCP6 instead of GCP3 by tagging Rab11 as bait (Fig. 6B). The co-IP results showed that CRB3 knockdown did not affect Rab11 interacting with GCP6, but the level of Rab11 binding to GCP6 tended to decrease in MCF10A cells with CRB3 knockdown (Fig. 6C). Consistent with other reports[30], Rab11 knockdown significantly inhibited primary cilium formation in MCF10A cells (Supplementary Fig. 5A-C).

Revised version: To fully prove this hypothesis, we first examined the interaction of CRB3 with Rab11-positive endosomes and γTuRC-specific proteins in MCF10A cells. Using CRB3 as bait, we identified that CRB3 could interact with Rab11, but not GCP6 and GCP3 (Fig. 6A). In contrast, Rab11 could interact with CRB3 and GCP6 instead of GCP3 by tagging Rab11 as bait (Fig. 6B). The co-IP results showed that CRB3 knockdown did not affect Rab11 interacting with GCP6, but the level of Rab11 binding to GCP6 tended to decrease in MCF10A cells with CRB3 knockdown (Fig. 6C). Consistent with other reports[30], Rab11 knockdown significantly inhibited primary cilium formation in MCF10A cells (Supplementary Fig. 5C, D).

CRB3 interacts with Rab11.

A. Coimmunoprecipitation of CRB3 with Rab11, GCP6 and GCP3 in MCF10A cells. B. Coimmunoprecipitation of Rab11 with CRB3, GCP6 and GCP3 in MCF10A cells. C. Coimmunoprecipitation of Rab11 with GCP6 in control and CRB3 knockdown MCF10A cells. D. Schematic diagram of CRB3b domains. E. Diagram truncations of CRB3b-GFP fusion proteins with serial C-terminal deletions. F. Domain mapping of CRB3b-GFP for Flag-Rab11a binding. Flag antibody co-IP of the full-length CRB3b-GFP and truncations of CRB3b-GFP with Flag-Rab11a were cotransfected into HEK293 cells for 48 h. Immunoblot analysis was performed using GFP and Flag antibodies. G. Coimmunoprecipitation of Rab11 mutant variants with full-length CRB3b-GFP. Flag antibody co-IP of the full-length CRB3b-GFP with Flag-Rab11aWT, Flag-Rab11a[Q70L], Flag-Rab11a[S20V] and Flag-Rab11a[S25N] were cotransfected into HEK293 cells for 48 h. Immunoblot analysis was performed using GFP and Flag antibodies. Bars represent the means ± SD, and the experiments were performed in triplicate; Unpaired Student’s t test, *P<0.05, ***P<0.001.

Based on these results, we further detected the region of CRB3b that interacts with Rab11 by using coexpression and co-IP in HEK293 cells. CRB3b is composed of a signal peptide (SP), extracellular domain, transmembrane (TM), FERM-binding domain (FBD), and carboxy-terminal PDZ-binding domain (PBD). We constructed CRB3-GFP fusion proteins by fusing GFP tags to the extracellular domain, as previously reported[31] (Fig. 6D). According to these domains, we also constructed the truncations of CRB3-GFP fusion proteins with serial C-terminal deletions (Fig. 6E) and another Flag-Rab11a fusion protein. The coexpression and co-IP results showed that Flag-Rab11a interacts with full-length CRB3 (1-120), CRB3 (1-116), CRB3 (1-83), and CRB3 (1-58) but not CRB3 (1-26) (Fig. 6F). Therefore, the region of CRB3b that interacts with Rab11 is amino acids 27 to 58.

Rab11 is a small GTP-binding protein that plays an essential role in regulating the dynamics of recycling endosomes, which must undergo GDP/GTP cycles. In particular, Rab11a[S25N] is a GDP-locked form, while Rab11a[S20V] and Rab11a[Q70L] are GTP-locked forms[30]. To further investigate the interaction between the GTP-bound form of Rab11a and CRB3b, we constructed these mutant variants of Rab11a. Next, we found that CRB3b interacts more strongly with Rab11a[Q70L] and Rab11a[S20V], while weakly interacting with Rab11a[S25N], compared to Rab11aWT (Fig. 6G). Together, these results indicate that the interaction of CRB3b and Rab11a is dependent on the GTPase activity of Rab11a.

CRB3 navigates GCP6/Rab11 trafficking vesicles to CEP290 in the primary cilium

Since CRB3 knockdown does not affect the interaction with Rab11-positive endosomes and γTuRC-specific proteins GCP6, our focus was to determine whether CRB3 affects the localization of GCP6/Rab11 trafficking vesicles to the basal body of the primary cilium. We reviewed the identification of CRB3 interacting proteins and found some centriolar proteins. CEP290, for example, is located in the transition zone of the primary cilium and is required for the formation of microtubule-membrane linkers[32]. Then, we verified that exogenous CRB3b could bind to CEP290, Rab11, GCP6 (Fig. 7A), and Rab11 also interacted with CEP290, CRB3-GFP, GCP6 (Fig. 7B). Similar to Fig. 6C, the amount of GCP6 bound to Rab11 was reduced within CRB3 knockdown. Importantly, CRB3 knockdown made it difficult for Rab11 to bind CEP290 (Fig. 7C). We corroborated these findings by checking the colocalization of the basal body foci of GCP6 and γ-tubulin, γ-tubulin and Rab11 in MCF10A and MEF cells. In quiescence control cells, GCP6 foci and γ-tubulin formed a prominent focus at the basal body of the primary cilium, as observed by immunofluorescence. However, CRB3 downregulation eliminated this focus and significantly disturbed the basal body foci of GCP6 and γ-tubulin in MCF10A cells (Fig. 7D), and γ-tubulin and Rab11 in MEF cells (Fig. 7E). At the beginning of this experiment, we also attempted to show the interaction of CRB3, Rab11 and GCP6 by expressing fluorescence tagged GCP6 protein and using FRET assay. However, the centrosomal foci of fluorescent GCP6 was difficult to capture due to weak signal, and this may require re-overexpression of ninein-like protein (Nlp) to promote γTuRC enrichment[29, 33]. Although CRB3 knockdown had little effect on the formation of GCP6/Rab11 trafficking vesicles, it significantly inhibited γTuRC assembly during ciliogenesis. This, in turn, impacted the transport of GCP6/Rab11 trafficking vesicles to the basal body of the primary cilium. Since CRB3 affected the binding of Rab11 and CEP290, we speculate that CRB3 may act as a navigator, navigating GCP6/Rab11 trafficking vesicles to the basal body of the primary cilium.

CRB3 navigates GCP6/Rab11 trafficking vesicles to the basal body of the primary cilium.

A. Coimmunoprecipitation of exogenous CRB3b with Rab11, GCP6, GCP3 and CEP290 in MCF10A cells. B. Coimmunoprecipitation of Rab11 with exogenous CRB3b, GCP6, GCP3 and CEP290 in MCF10A cells. C. Coimmunoprecipitation of Rab11 with GCP6 and CEP290 in control and CRB3 knockdown MCF10A cells. D. Representative immunofluorescent images of GCP6 and γ-tubulin colocalization of the basal body foci in MCF10A cells with CRB3 knockdown. γ-tubulin (green), GCP6 (red), and DNA (blue). (Colocalization is marked by arrows; scale bars, 25 μm) E. Representative immunofluorescent images of Rab11 and γ-tubulin colocalization of the basal body foci in MEF cells from Crb3fl/fl and Crb3fl/fl;CAG-Cre mice. γ-tubulin (green), Rab11 (red), and DNA (blue). Foci are marked by arrows; scale bars, 25 μm. Bars represent the means ± SD; Unpaired Student’s t test, ***P<0.001.

CRB3 regulates the Hh and Wnt signaling pathways in tumorigenesis

To investigate the role of CRB3 in regulating ciliary assembly in breast cancer, we examined the relationship between CRB3 and the primary cilium in breast cancer tissues. We found that CRB3 was localized at the subapical surface of the mammary gland lumen in the paracarcinoma tissues, but no obvious expression or localization in breast cancer tissues (Fig. 8A). Immunofluorescence analysis showed that the intact primary cilium was observed in the adjacent para-carcinoma tissues, whereas the number of cells with primary cilia was much lower in breast cancer tissues (Fig. 8B, C). Consistent with previous findings, CRB3 was localized at the basal body of the primary cilium in adjacent para-carcinoma tissues (Fig. 8D). These results again suggest that a defect in CRB3 expression inhibits ciliary assembly, which could be involved in tumorigenesis.

Defects in CRB3 expression inhibit ciliary assembly in breast cancer tissues, and activate the Wnt signaling pathway in mammary cells and PyMT mouse model.

A. Representative immunofluorescent images of CRB3 in breast cancer tissues (n=50). CRB3 (green) and DNA (blue). (scale bars, 25 μm) B. Representative immunofluorescent images of the primary cilium in breast cancer tissues (n=50). Acetylated tubulin (red), γ-tubulin (green), and DNA (blue). (scale bars, 10 μm) C. Quantification of the proportion of cells with primary cilium formation in breast cancer tissues (n=50). D. Representative immunofluorescent images of CRB3 and γ-tubulin colocalization in adjacent paracarcinoma tissues. CRB3 (green), γ-tubulin (red), and DNA (blue). (scale bars, 25 μm) E. Quantification of the proportion of MFC10A cells with SMO translocation after CRB3 knockdown (n=6). F. Real-time quantitative PCR showing the relative mRNA expression of GLI1 upon SAG treatment in CRB3-depleted MFC10A cells (n=6). G. Immunoblot analyses of the effect of CRB3 on the molecules of the Wnt signaling pathway in mammary cells, and the experiments were performed in triplicate. H. I. Immunohistochemical analyses of GLI1 and β-catenin in primary tumors from PyMT-WT and PyMT-cKO-Crb3 mice at 9 weeks old, respectively. (scale bars, 25 μm) Bars represent the means ± SD; Unpaired Student’s t test, ***P<0.001.

Earlier version: To investigate the effects of primary cilium absence, we examined alterations in the Hedgehog (Hh) and Wnt signaling pathways in CRB3-depleted cells. Firstly, we detected the translocation of SMO into the primary cilium and the expression of the target gene GLI1. SAG, a potent Smoothened (SMO) receptor agonist, activates the Hh signaling pathway. However, in MCF10A cells treated with SAG, CRB3 knockdown did not promote SMO translocation into the primary cilium (Fig 8E). After SAG treatment, the mRNA expression of GLI1 was significantly decreased in CRB3-depleted cells (Fig 8F). Crb3 knockout did not alter GLI1 expression and subcellular localization in primary tumors of PyMT-cKO-Crb3 mice (Fig 8H). We previously reported that CRB3 expression was high in immortalized mammary epithelial cells, whereas loss of CRB3 expression was observed in breast cancer cells[6]. On the other hand, western blotting showed that CRB3 knockdown resulted in significant downregulation of GSK3-β in MCF10A cells. Additionally, β-catenin, the major effector of the canonical Wnt signaling pathway, was upregulated. In breast cancer cells, CRB3b overexpression led to upregulation of GSK3-β and downregulation of β-catenin (Fig. 8G, Supplementary Fig. 6). Notably, β-catenin was significantly upregulated and markedly localized in the nucleus of PyMT-cKO-Crb3 mice (Fig 8I). Overall, these data suggest that CRB3 knockdown cannot activate the Hh signaling pathway but can activate the Wnt signaling pathway, which is ultimately contributes to tumorigenesis.

Revised version: To investigate the effects of primary cilium absence, we examined alterations in the Hedgehog (Hh) and Wnt signaling pathways in CRB3-depleted cells. Firstly, we detected the translocation of SMO into the primary cilium and the expression of the target gene GLI1. SAG, a potent Smoothened (SMO) receptor agonist, activates the Hh signaling pathway. However, in MCF10A cells treated with SAG, CRB3 knockdown did not promote SMO translocation into the primary cilium (Fig 8E). After SAG treatment, the mRNA expression of GLI1 was significantly decreased in CRB3-depleted cells (Fig 8F). Crb3 knockout did not alter GLI1 expression and nuclear localization in primary tumors of PyMT-cKO-Crb3 mice (Fig 8H). We previously reported that CRB3 expression was high in immortalized mammary epithelial cells, whereas loss of CRB3 expression was observed in breast cancer cells[6]. On the other hand, western blotting showed that CRB3 knockdown resulted in significant downregulation of GSK3-β in MCF10A cells. Additionally, β-catenin, the major effector of the canonical Wnt signaling pathway, was upregulated. In breast cancer cells, CRB3b overexpression led to upregulation of GSK3-β and downregulation of β-catenin (Fig. 8G, Supplementary Fig. 6). Notably, β-catenin was significantly upregulated and markedly localized in the nucleus of PyMT-cKO-Crb3 mice (Fig 8I). Overall, these data suggest that CRB3 knockdown cannot activate the Hh signaling pathway but can activate the Wnt signaling pathway, which is ultimately contributes to tumorigenesis.

Discussion

Earlier version: Apical-basal cell polarity is essential for cellular architecture, development and homeostasis in epithelial tissues. However, most malignant cancers lack cell polarity due to unconstrained cell cycle and cell migration with decreased adhesion. CRB3, an important apical protein, is significantly reduced in various tumor cells and tissues, and promotes metastasis and tumor formation in nude mice[5, 6, 8, 34–36]. However, it has also been reported that CRB3 sequencing results are upregulated in breast cancer samples[37]. Here, we generated a novel transgenic mouse model with conditional deletion of Crb3 to investigate the role of polarity proteins in tumorigenesis. Crb3 knockout mice were reported to have cystic kidneys, improper airway clearance in the lung, villus fusion, apical membrane blebs, and disrupted microvilli in intestine airway clearances, which demonstrated that CRB3 is crucial for the development of the apical membrane and epithelial morphogenesis[19, 38]. Importantly, our studies in the Crb3 knockout mouse model suggest that, in addition to death after birth, it is characterized by smaller size, ocular abnormality, bronchial smooth muscle layer defect, reduced pancreatic islets and hyperplasia mammary ductal and renal tubular epithelium (some data not shown). These important phenotypic changes are due to disruption of epithelial polarity homeostasis, which may ultimately lead to tumorigenesis.

Revised version: Apical-basal cell polarity is essential for cellular architecture, development, and homeostasis in epithelial tissues. However, most malignant cancers lack cell polarity due to the unconstrained cell cycle and cell migration with decreased adhesion. CRB3, an important apical protein, is significantly reduced in various tumor cells and tissues, as well as promotes metastasis and tumor formation in nude mice[5, 6, 8, 34–36]. Here, we generated a novel transgenic mouse model with conditional deletion of Crb3 to investigate the role of polarity proteins in tumorigenesis. Crb3 knockout mice were reported to have cystic kidneys, improper airway clearance in the lung, villus fusion, apical membrane blebs, and disrupted microvilli in intestine airway clearances, which demonstrated that CRB3 is crucial for the development of the apical membrane and epithelial morphogenesis[19, 37]. Importantly, our studies in the Crb3 knockout mouse model suggest that in addition to death after birth, it is characterized by smaller size, ocular abnormality, bronchial smooth muscle layer defect, reduced pancreatic islets, and hyperplasia mammary ductal and renal tubular epithelium (some data not shown). These important phenotypic changes are due to the disruption of epithelial polarity homeostasis, which may ultimately lead to tumorigenesis.

Earlier version: In addition, we observed the absence of primary cilia in the mammary ductal lumen and renal tubule in the Crb3 knockout mouse model. Consistent with other reports, CRB3 knockdown regulated ciliary assembly in MDCK cells[17, 18]. Furthermore, CRB3 was localized in the inner segments of photoreceptor cells and concentrated with the connecting cilium throughout the development of the mouse retina[39]. These ocular defect phenotypes closely resemble those of EHD1 knockout mice, which regulate ciliary vesicle formation in primary cilium assembly[25, 40, 41]. Hence, CRB3 can affect primary cilium formation in various epithelial tissues, such as the breast, kidney, and retina.

Revised version: In addition, we observed the absence of primary cilia in the mammary ductal lumen and renal tubule in the Crb3 knockout mouse model. Consistent with other reports, CRB3 knockdown regulated ciliary assembly in MDCK cells[17, 18]. Furthermore, CRB3 was localized in the inner segments of photoreceptor cells and concentrated with the connecting cilium throughout the development of the mouse retina[38]. These ocular defect phenotypes closely resemble those of EHD1 knockout mice, which regulate ciliary vesicle formation in primary cilium assembly[25, 39, 40]. Hence, CRB3 can affect primary cilium formation in various epithelial tissues, such as the breast, kidney, and retina.

Earlier version: Previous literature has reported that CRB3 localizes in the primary cilium itself in differentiated MDCK cells[17, 18], but we mainly found that it is not located on the basal body of the primary cilium in mammary epithelial cells. The immunofluorescence results in these literatures showed that CRB3 was scattered on the primary cilia, with a strong foci at the basal body. Notably, in rat kidney collecting ducts, the localization of CRB3 on primary cilia was significantly reduced, while the obvious localization was at the basal body[18]. Another study also reported the co-localization of CRB3b and γ-tubulin in MDCK cells[17]. This finding is consistent with our conclusion, and we also verified its co-localization with the centrosome by overexpressing CRB3b in mammary epithelial cells, indicating that CRB3 mainly localizes to the basal body of the primary cilium. Since CRB3b also alters primary cilium formation in the mammary epithelium, it is important to investigate the molecular mechanisms by which CRB3 regulates primary cilium formation.

Revised version: Previous literature has reported that CRB3 localizes within the primary cilium in differentiated MDCK cells[17, 18]. However, our findings indicate that it is mainly located on the basal body of the primary cilium in mammary epithelial cells. The immunofluorescence results in these literatures showed that CRB3 was scattered on the primary cilia, with a strong foci at the basal body. Notably, in rat kidney collecting ducts, the localization of CRB3 on primary cilia was significantly reduced, while the obvious localization was at the basal body[18]. Another study also reported the co-localization of CRB3b and γ[tubulin in MDCK cells[17]. This finding is consistent with our conclusion, and we also verified its co-localization with the centrosome by overexpressing CRB3b in mammary epithelial cells, indicating that CRB3 mainly localizes to the basal body of the primary cilium. Since CRB3b also alters primary cilium formation in the mammary epithelium, it is important to investigate the molecular mechanisms by which CRB3 regulates primary cilium formation. Subsequently, we used mass spectrometry to identify the CRB3b binding protein. Our findings indicate that CRB3b not only binds to some centrosomal proteins, but also plays a significant role in Golgi vesicle transport (Figure 5A, B). These results are consistent with the results of the endogenous CRB3 staining by laser confocal microscopy (Figure 4C). Specifically, endogenous CRB3 localizes to the centrosome and strongly accumulates in an undefined intracellular structure near the cell nucleus, likely to be the Golgi. However, when we compared the fluorescence microscopy analysis of exogenous CRB3b staining (Figure 4A), we found that the main localization was observed at the cell tight junctions, which is consistent with the reported molecular function. Additionally, there was a strong accumulation of staining around the cell nucleus, with a shape different from the endogenous staining around the cell nucleus. We speculate that this may be due to excessive overexpression of exogenous proteins, which obstruct the protein secretion pathway and result in the accumulation of the exogenous protein at the ER and Golgi. To better study the cell localization of exogenous CRB3b, a conditional induction expression or knock-in tag on the endogenous CRB3 gene could be used.

Earlier version: Polarized vesicular traffic plays a crucial role in extending the primary cilium on apical epithelial surfaces. However, the molecular mechanism and its relationship to tumorigenesis remain unclear. Here, we reveal that Rab11-positive endosomes mediate the intracellular trafficking of CRB3 and that CRB3 can navigate GCP6/Rab11 trafficking vesicles to CEP290, promoting correct γTuRC assembly during primary cilium formation. When cells form tight junctions, CRB3 is mainly localized to the cell membrane, and subsequently the endosomes participate in the intracellular degradation process of CRB3 on the cell membrane. The intracellular CRB3 can bind to Rab11 through the extracellular domain, which in turn participates in primary cilium assembly. Additionally, in the presence of CRB3, the epithelium can form contact inhibition and an intact primary cilium in quiescence to achieve cellular homeostasis. In contrast, the loss of contact inhibition and primary cilium formation with CRB3 deletion can activate the Wnt signaling pathway, affect the Hh signaling pathway, and disrupt the imbalance of this cellular homeostasis, resulting in significant cell proliferation during tumorigenesis (Fig. 9).

Revised version: Polarized vesicular traffic plays a crucial role in extending the primary cilium on apical epithelial surfaces. However, the molecular mechanism and its relationship to tumorigenesis remain unclear. Here, we reveal that Rab11-positive endosomes mediate the intracellular trafficking of CRB3 and that CRB3 can navigate GCP6/Rab11 trafficking vesicles to CEP290, promoting correct γTuRC assembly during primary cilium formation. When cells form tight junctions, CRB3 is mainly localized to the cell membrane, and subsequently, the endosomes participate in the intracellular degradation process of CRB3 on the cell membrane. The intracellular CRB3 can bind to Rab11, which in turn participates in primary cilium assembly. Additionally, in the presence of CRB3, the epithelium can form contact inhibition and an intact primary cilium in quiescence to achieve cellular homeostasis. In contrast, the loss of contact inhibition and primary cilium formation with CRB3 deletion can activate the Wnt signaling pathway, affect the Hh signaling pathway, and disrupt the imbalance of this cellular homeostasis, resulting in significant cell proliferation during tumorigenesis (Fig. 9).

Schematic model of CRB3 regulating ciliary assembly.

Earlier version: Graphic summary of prominent phenotypes observed after CRB3 deletion. CRB3 is localized on apical epithelial surfaces and participates in tight junction formation to maintain contact inhibition and cell homeostasis in quiescence. Inside the cell, Rab11-positive endosomes mediate the intracellular trafficking of CRB3, and CRB3navigates GCP6/Rab11 trafficking vesicles to CEP290. Then, GCP6 is involved in normal γTuRC assembly in ciliogenesis. In CRB3 deletion cells, the primary cilium fails to assemble properly, and the Wnt signaling pathway is activated through β-catenin upregulation and nuclear localization, but the Hh signaling pathway fails to be activated. This cellular imbalance is disrupted, leading to tumorigenesis. Abbreviations: γTuRC, γ-tubulin ring complex; TJ, tight junction; EE, early endosome; LE, late endosome.

Revised version: Graphic summary of prominent phenotypes observed after CRB3 deletion. CRB3 is localized on apical epithelial surfaces and participates in tight junction formation to maintain contact inhibition and cell homeostasis in quiescence. Inside the cell, Rab11-positive endosomes mediate the intracellular trafficking of CRB3, and CRB3 navigates GCP6/Rab11 trafficking vesicles to CEP290. Then, GCP6 is involved in normal γTuRC assembly in ciliogenesis. In CRB3 deletion cells, the primary cilium fails to assemble properly, and the Wnt signaling pathway is activated through β-catenin upregulation and nuclear localization, but the Hh signaling pathway fails to be activated. This cellular imbalance is disrupted, leading to tumorigenesis. Abbreviations: γTuRC, γ-tubulin ring complex; TJ, tight junction; EE, early endosome; LE, late endosome.

Earlier version: The formation of the primary cilium has been divided into several distinct phases. It begins with the maturation of the mother centriole, also known as the centriole-to-basal-body transition[9, 42]. DAVs, which are small cytoplasmic vesicles originating from the Golgi and the recycling endosome, accumulate in the vicinity of the distal appendages of the mother centriole[25, 43, 44]. In growing RPE1 cells, Rab8a is localized to cytoplasmic vesicles and the Golgi/trans-Golgi network, while Rab8-positive vesicles are rapidly redistributed to the mother centriole by binding to Rab11-positive recycling and endosomes under serum starvation[26, 45]. The EHD1 protein converts these DAVs to form larger ciliary vesicles, which elongate through continuous fusion with Rab8-positive vesicles to produce the primary cilium membrane[25, 30, 41]. Therefore, Rab11-positive recycling endosomes, as important transporters, are necessary for primary cilium formation and centriole-basal-body transition in the early phase. Our study found that CRB3 could bind to some members of Rab small GTPase family, significantly participating in Golgi vesicle transport and vesicle organization pathways through mass spectrometry identification. And coexpression and co-IP results showed that CRB3 interacted with Rab11 via the extracellular domain, which mediated the intracellular trafficking of CRB3 within Rab11-positive endosomes at different cell densities.

Revised version: The formation of the primary cilium has been divided into several distinct phases. It begins with the maturation of the mother centriole, also known as the centriole-to-basal-body transition[9, 41]. DAVs, which are small cytoplasmic vesicles originating from the Golgi and the recycling endosome, accumulate in the vicinity of the distal appendages of the mother centriole[25, 42, 43]. In growing RPE1 cells, Rab8a is localized to cytoplasmic vesicles and the Golgi/trans-Golgi network, while Rab8-positive vesicles are rapidly redistributed to the mother centriole by binding to Rab11-positive recycling and endosomes under serum starvation[26, 44]. The EHD1 protein converts these DAVs to form larger ciliary vesicles, which elongate through continuous fusion with Rab8-positive vesicles to produce the primary cilium membrane[25, 30, 40]. Therefore, Rab11-positive recycling endosomes, as important transporters, are necessary for primary cilium formation and centriole-basal-body transition in the early phase. Our study found that CRB3 could bind to some members of Rab small GTPase family, significantly participating in Golgi vesicle transport and vesicle organization pathways through mass spectrometry identification. And coexpression and co-IP results showed that CRB3 interacted with Rab11, which mediated the intracellular trafficking of CRB3 within Rab11-positive endosomes at different cell densities.

Earlier version: Unlike motile cilia, primary cilia have a 9+0 structure, consisting of only nine outer doublet microtubules without a central pair. Similar to microtubule assembly at microtubule organizing centers (MTOCs) during mitosis, γ-tubulin, which is a part of γTuRC, is also concentrated in the basal body of the cilium in interphase[46]. γ-tubulin complexes have been shown to form microtubule templates and regulate microtubule nucleation through longitudinal contacts with α[tubulin and β[tubulin[27]. The γTuRC consists of multiple copies of γTuSC, which is composed of two copies of γ-tubulin and one each of GCP2 and GCP3. Then, GCP2 and GCP3 interact with GCP4, GCP5 and GCP6 to form γTuRC[47]. Our experiments showed that CRB3 knockdown did not affect the expression of the γTuRC-specific proteins GCP3, GCP6 and γ-tubulin, but inhibited the interaction between GCP6 and GCP3, leading to destabilization of γTuRC. Further analysis revealed that significant colocalization was observed between GCP6- and Rab11-positive vesicles at low cell density, while at high cell density, GCP6 accumulated at the centrosome with GCP3. CRB3 knockdown resulted in the disappearance of this accumulated GCP6. These results indicated that GCP6/Rab11 trafficking vesicles must be transported to the basal body for ciliogenesis. While it is still unclear how CRB3 interferes with γTuRC assembly during primary cilium formation, we hypothesize that Rab11-positive endosomes mediate the intracellular trafficking of CRB3 and that CRB3, as a navigator interacting with CEP290, can navigate GCP6/Rab11 trafficking vesicles to the transition zone for γTuRC assembly during primary cilium formation.

Revised version: Unlike motile cilia, primary cilia have a 9+0 structure, consisting of only nine outer doublet microtubules without a central pair. Similar to microtubule assembly at microtubule[organizing centers (MTOCs) during mitosis, γ-tubulin, which is a part of γTuRC, is also concentrated in the basal body of the cilium in interphase[45]. γ[tubulin complexes have been shown to form microtubule templates and regulate microtubule nucleation through longitudinal contacts with α[tubulin and β[tubulin[27]. The γTuRC consists of multiple copies of γTuSC, which is composed of two copies of γ[tubulin and one each of GCP2 and GCP3. Then, GCP2 and GCP3 interact with GCP4, GCP5 and GCP6 to form γTuRC[46]. Our experiments showed that CRB3 knockdown did not affect the expression of the γTuRC-specific proteins GCP3, GCP6, and γ[tubulin, but inhibited the interaction between GCP6 and GCP3, leading to the destabilization of γTuRC. Further analysis revealed that significant colocalization was observed between GCP6- and Rab11-positive vesicles at low cell density, while at high cell density, GCP6 accumulated at the centrosome with GCP3. CRB3 knockdown resulted in the disappearance of this accumulated GCP6. These results indicated that GCP6/Rab11 trafficking vesicles must be transported to the basal body for ciliogenesis. While it is still unclear how CRB3 interferes with γTuRC assembly during primary cilium formation, we hypothesize that Rab11-positive endosomes mediate the intracellular trafficking of CRB3 and that CRB3, as a navigator interacting with CEP290, can navigate GCP6/Rab11 trafficking vesicles to the transition zone for γTuRC assembly during primary cilium formation.

Earlier version: CEP290 is a centriolar or ciliary protein that localizes to the centrosomes in dividing cells, the distal mother centriole in quiescent cells, and the transition zone in the primary cilium[48]. Mutations in CEP290 are associated with various diseases, such as the devastating blinding disease Leber’s Congenital Amaurosis (LCA), nephronophthisis, Senior Lǿken syndrome (SLS), Joubert syndrome (JS), Bardet-Biedl syndrome (BBS), and lethal Meckel-Gruber syndrome (MGS)[49]. In particular, the phenotypes of CEP290-associated ciliopathies are very similar to those in our Crb3 knockout mouse model, with lethality and ocular and renal abnormalities. CEP290 knockdown significantly decreased the number of cells with primary cilia, and CEP290 knockout mice showed mislocalization of ciliary proteins in the retinas[50, 51]. Additionally, depletion of CEP290 disrupts protein trafficking to centrosomes and affects the recruitment of the BBSome and vesicular trafficking in ciliary assembly[52, 53]. Although CEP290 serves as a hub to recruit many ciliary proteins, it has not been reported to be involved in polarized vesicle trafficking or localization. Our study verified that CEP290 could interact with the polarity protein CRB3. CRB3 knockdown disrupts the interaction between CEP290 and Rab11-positive endosomes. Thus, these results validate our hypothesis. Although CRB3 navigates GCP6/Rab11 trafficking vesicles to CEP290, we further need to detect the region of CRB3 interacting with CEP290 and what other cargos are transported by CRB3/Rab11 trafficking vesicles in ciliary assembly.

Revised version: CEP290 is a centriolar or ciliary protein that localizes to the centrosomes in dividing cells, the distal mother centriole in quiescent cells, and the transition zone in the primary cilium[47]. Mutations in CEP290 are associated with various diseases, such as the devastating blinding disease Leber’s Congenital Amaurosis (LCA), nephronophthisis, Senior Lǿken syndrome (SLS), Joubert syndrome (JS), Bardet-Biedl syndrome (BBS), and lethal Meckel-Gruber syndrome (MGS)[48]. In particular, the phenotypes of CEP290-associated ciliopathies are very similar to those in our Crb3 knockout mouse model, with lethality and ocular and renal abnormalities. CEP290 knockdown significantly decreased the number of cells with primary cilia, and CEP290 knockout mice showed mislocalization of ciliary proteins in the retinas[49, 50]. Additionally, the depletion of CEP290 disrupts protein trafficking to centrosomes and affects the recruitment of the BBSome and vesicular trafficking in ciliary assembly[51, 52]. Although CEP290 serves as a hub to recruit many ciliary proteins, it has not been reported to be involved in polarized vesicle trafficking or localization. Our study verified that CEP290 could interact with the polarity protein CRB3. CRB3 knockdown disrupts the interaction between CEP290 and Rab11-positive endosomes. Thus, these results validate our hypothesis. Although CRB3 navigates GCP6/Rab11 trafficking vesicles to CEP290, we further need to detect the region of CRB3 interacting with CEP290 and what other cargos are transported by CRB3/Rab11 trafficking vesicles in ciliary assembly.

Earlier version: In addition, our results showed that CRB3 knockdown failed to activate the Hh signaling pathway with SAG in MCF10A cells. Under untreated conditions, PTCH1 located at the ciliary membrane inhibits SMO, and Suppressor of Fused (SuFu) sequesters Gli transcription factors, leading to their degradation. SAG can activate the Hh signaling pathway, resulting in dissociation into the nucleus to regulate the expression of downstream target genes[54]. The Hh signaling pathway mainly regulates cellular growth and differentiation and is abnormal in different types of tumors. Some studies indicate that cilia are double-sided, promoting and preventing cancer development through the Hh signaling pathway in vivo. It is thought that cilia deletion could activate SMO to inhibit tumor growth, while promoting carcinogenesis was induced by activated Gli2[55]. Although SAG cannot activate the Hh signaling pathway in CRB3-depleted cells, further research is needed to determine the role of CRB3 in the carcinogenic process through the ciliary Hh signaling pathway. Additionally, another canonical Wnt signaling pathway was altered after CRB3 knockdown. The Wnt signaling pathway is an important cascade regulating development and stemness, and aberrant Wnt signaling has been reported in many cancer entities, especially colorectal cancer[56]. In our study, CRB3 downregulated the expression of β-catenin. Thus, CRB3 can affect cilium-related signaling pathways in tumorigenesis.

Revised version: In addition, our results showed that CRB3 knockdown failed to activate the Hh signaling pathway with SAG in MCF10A cells. Under untreated conditions, PTCH1 located at the ciliary membrane inhibits SMO, and Suppressor of Fused (SuFu) sequesters Gli transcription factors, leading to their degradation. SAG can activate the Hh signaling pathway, resulting in dissociation into the nucleus to regulate the expression of downstream target genes[53]. The Hh signaling pathway mainly regulates cellular growth and differentiation and is abnormal in different types of tumors. Some studies indicate that cilia are double-sided, promoting and preventing cancer development through the Hh signaling pathway in vivo. It is thought that cilia deletion could activate SMO to inhibit tumor growth, while promoting carcinogenesis was induced by activated Gli2[54]. Although SAG cannot activate the Hh signaling pathway in CRB3-depleted cells, further research is needed to determine the role of CRB3 in the carcinogenic process through the ciliary Hh signaling pathway in PyMT-cKO-Crb3 mice. Additionally, another canonical Wnt signaling pathway was altered after CRB3 knockdown. The Wnt signaling pathway is an important cascade regulating development and stemness, and aberrant Wnt signaling has been reported in many cancer entities, especially colorectal cancer[55]. In our study, CRB3 downregulated the expression of β-catenin. Thus, CRB3 can affect cilium-related signaling pathways in tumorigenesis.

In summary, our study provides novel evidence for the role of the apical polarity protein CRB3 in regulating viability, mammary and ocular development, and primary cilium formation in germline and conditional knockout mice. Deletion of CRB3 caused irregular lumen formation and ductal epithelial hyperplasia in mammary epithelial-specific knockout mice. Further study uncovered a new function of CRB3b in interacting with Rab11, navigating GCP6/Rab11 trafficking vesicles to CEP290 for γTuRC assembly during ciliogenesis. In addition, CRB3 also participates in regulating cilium-related Hh and Wnt signaling pathways in tumorigenesis. Understanding these polarity protein-mediated vesicle trafficking mechanisms will shed new light on luminal development and tumorigenesis.

Methods

Generation of floxed Crb3 mice

All animal experiments were verified and approved by the Committee of Institutional Animal Care and Use of Xi’an Jiaotong University. Heterozygous Crb3wt/fl mice (C57BL/6J) were generated using Cas9/CRISPR-mediated genome editing (Cyagen Biosciences, Guangzhou, China). The gRNA to Crb3 and Cas9 mRNA was coinjected into fertilized mouse eggs to generate targeted conditional knockout offspring. The Crb3 gRNA sequences were gRNA1, 5’-GGCTGGGTCCACACCTACGGAGG-3’; gRNA2, 5’-ACCCACAAAGCCACGCCA GTGGG-3’. Mouse genotyping was screened using PCR with the primers F1, 5’-ACATAAGGCCTTCCGTTAAGCTG-3’; R1, 5’-GTGGATTCGGACCAGTCTGA-3’. Crb3wt/fl mice were intercrossed with MMTV-Cre mice or CAG-Cre mice to generate tissue-specific Crb3 knockout mice. MMTV-Cre mice (FVB), CAG-Cre mice (C57BL/6N), and MMTV-PyMT mice (FVB) were purchased from Cyagen Biosciences. All mice were bred in the specific pathogen-free (SPF) animal houses of the Laboratory Animal Center of Xi’an Jiaotong University.

Mice were genotyped using PCR and DNA gel electrophoresis. Genomic DNA from the mouse tail was extracted using the Mouse Direct PCR Kit (Bimake, TX, USA, #B40015) according to the manufacturer’s instructions. The genotyping primers were F3, 5’-TTGAGAGTCTTAAGCAGTCAGGG-3’; R5, 5’- AACCTTTCCCAGGAGTA TGTGAC-3’. PCR results showed that Crb3wt/wt was one band with 163 bp, Crb3wt/fl was two bands with 228 bp and 163 bp, and Crb3fl/fl was one band with 228 bp. The primers for identifying Cre alleles were F, 5’-TTGAGAGTCTTAAGCAGTCAGGG-3’; R, 5’-TTACCACTCCCAGCAAGACAC-3’, and amplification with one 247 bp band. The reaction conditions of Crb3 and Cre PCR were as follows: 94°C for 5 min, 94°C for 20 s, 60°C for 30 s, and 72°C for 20 s for 35 cycles.

Mouse mammary gland whole mount analysis

After euthanizing the mice, the inguinal mammary gland tissues were removed intact and fully spread onto slides. The tissues were rapidly fixed in Carnoy’s fixative for 2 h at room temperature. Subsequently, the tissues were washed in 70% ethanol solution for 15 min at room temperature, and in 50% and 30% ethanol and distilled water for 5 min each. Staining was performed in carmine alum solution at 4°C overnight. After washing with a 70% ethanol solution, the tissues were washed again with 70%, 95%, and 100% ethanol for 15 min each. Tissues were cleared in xylene overnight at room temperature and then mounted with neutral balsam. Images were obtained using a microscope (Leica DMi8, Wetzlar, Germany).

Cell culture and transfection

The MCF10A, MCF7, T47D, MDA-MB-231, HCC1806, T47D, and MDA-MB-453 cell lines were purchased from Shanghai Institute of Biochemistry and Cell Biology (Chinese Academy of Sciences, Shanghai, China). MCF10A cells were routinely grown in DMEM/F12 (1:1) media (HyClone, UT, USA) supplemented with 5% horse serum, 20 ng/ml human EGF, 10 μg/ml insulin, and 0.5 μg/ml hydrocortisone. MCF7, MDA-MB-231, and MDA-MB-453 cells were cultured in high glucose DMEM media (HyClone, UT, USA) supplemented with 10% FBS (HyClone, UT, USA, #SH30084.03). T47D and HCC1806 cells were grown in RPMI-1640 media (HyClone, UT, USA) supplemented with 10% FBS. All cell lines were incubated in 5% CO2 at 37°C.

siRNAs were purchased from GenePharma company (Shanghai, China). Cultured cells were transfected with siRNAs or plasmids using Lipofectamine 2000 (Invitrogen, CA, USA, #11668-019) following the manufacturer’s protocol. The negative control shRNA (NC) and shRNA against CRB3 (both isoforms CRB3a and CRB3b) were packaged into lentivirus as previously described[5, 6].

DNA constructs and stable cell lines

CRB3b was PCR amplified from RNA extracted from MCF10A cells and cloned into the pCMV-Blank vector (Beyotime, Shanghai, China, D2602) using T4 ligase (NEB, MA, USA). Then, the GFP sequence was inserted between the extracellular and transmembrane CRB3 domains to generate the pCMV-CRB3-GFP vector, as described previously[31]. The full-length CRB3 (1-120), CRB3 (1-116), CRB3 (1-83), CRB3 (1-58), and CRB3 (1-26) were amplified by PCR from the pCMV-CRB3-GFP vector using specific primers to create EcoR [and Xho [restriction sites, and then cloned into the pCMV-Blank vector. The full-length was also cloned into the lentiviral vector pLVX-TetOne-Puro (Clontech, Takara, #631849). The plasmid pECMV-3×FLAG-RAB11A was purchased from SinoBiological (Beijing, China). The plasmids of pECMV-3×FLAG-RAB11A[S20V]/[S25N]/[Q70L] mutants were constructed using the Fast Mutagenesis System (Transgen, Beijing, China, FM111-01).

Earlier version: Lentiviruses were produced in HEK293T cells. CRB3-GFP in MCF10A and MCF7 cells was generated using a lentivirus expression system. Stable cell lines were screened using puromycin after lentiviral infection for 72 h.

Revised version: Lentiviruses were produced in HEK293T cells. CRB3-GFP in MCF10A and MCF7 cells were generated using a lentivirus expression system. Stable cell lines were screened using puromycin after lentiviral infection for 72 h.

3D morphogenesis

To initiate 3D morphogenesis, growth factor-reduced Matrigel (Corning, USA, #354230) was added to a four-well chamber slide system (Corning, USA, #177437) and then incubated at 37°C for solidification. MCF10A cells were plated into this chamber slide and cultured in DMEM/F12 (1:1) media (HyClone, UT, USA) supplemented with 2% horse serum, 5 ng/ml human EGF, 10 μg/ml insulin, 0.5 μg/ml hydrocortisone, and 2.5% Matrigel. 3D morphogenesis was photographed using a microscope (Leica DMi8, Wetzlar, Germany) at different time points. On day 14 after cell culture, 3D MCF10A cells were fixed and stained.

Immunohistochemistry

The paraffin-embedded sections of samples were baked, deparaffinized in xylene, and sequentially rehydrated in gradient concentrations of ethanol, sequentially. Then, antigen retrieval was performed in Tris-EDTA buffer (pH 9.0) heated to 95°C for 20 min. Endogenous peroxidase activity was reduced in 3% hydrogen peroxide for 10 min at room temperature. Then, 5% goat plasma was used for blocking at 37°C for 30 min, and the sections were incubated with specific primary antibodies overnight at 4°C. Primary antibodies used were Ki67 (Abcam, UK, ab279653), phospho-histone H3 (Abcam, UK, ab267372), cleaved caspase 3 (CST, USA, #9661), GLI1 (CST, USA, #3538), and β-catenin (CST, USA, #8480). The following steps used biotin-streptavidin HRP detection kits (ZSGB-BIO, China, SP-9001, SP-9002) according to the manufacturer’s instructions. The sections were stained with DAB (ZSGB-BIO, China, ZLI-9017), counterstained with hematoxylin, dehydrated in gradient concentrations of ethanol, cleared in xylene, and mounted with neutral balsam, sequentially. Images were obtained using a slide scanner (Leica SCN400, Wetzlar, Germany).

Immunofluorescence

Cells were plated on coverslips and fixed in 4% paraformaldehyde for 30 min. Then, the membranes were permeabilized using 0.2% Triton X-100. The cells were blocked with 5% BSA solution for 1 h at room temperature. Cells were incubated with specific primary antibodies overnight at 4°C. Caspase 3 (CST, USA, #9662S), α-tubulin (CST, USA, #3873), acetylated tubulin (Proteintech, USA, #662001-IG), γ-tubulin (Proteintech, USA, #15176-1-AP), pericentrin (Abcam, UK, ab4448), CRB3 (Sigma, USA, HPA013835), EEA1 (BD, USA, #610456), Rab11 (BD, USA, #610823), and GCP6 (Abcam, UK, ab95172) were used as primary antibodies. The secondary antibodies were Alexa Fluor 488-labeled or 594-labeled (Invitrogen, USA, A32731, A32744). DPIA (5 μg/ml) was used for DNA staining, and images were captured using a confocal microscope (Leica SP5 [, Wetzlar, Germany).

MEFs isolation and maintenance

Earlier version: The embryos were harvested from Crb3wt/fl;CAG-Cre mice crosses at E13.5, and the head, tail, and internal organs were removed. The tail was used for genotyping, and DNA was extracted from it. Then, the body was minced and digested with trypsin in a 37°C incubator for 30 min. The resulting cells were resuspended in MEF media (DME, 10% FBS, 1% penicillin/streptomycin) and plated onto 10 cm cell culture dishes with one embryo in each. They were subsequently passaged 1:3 with MEF media to expand the culture.

Revised version: The embryos were harvested from Crb3wt/fl;CAG-Cre mice cross at E13.5, and the head, tail, and internal organs were removed. The tail was used for genotyping, and DNA was extracted from it. Then, the body was minced and digested with trypsin in a 37°C incubator for 30 min. The resulting cells were resuspended in MEF media (DME, 10% FBS, 1% penicillin/streptomycin) and plated onto 10 cm cell culture dishes with one embryo in each. They were subsequently passaged 1:3 with MEF media to expand the culture.

Real-time PCR assay

Total RNA from cell lines was isolated using TRIzol reagent (Invitrogen, USA, #15596026), and 5 μg RNA was converted to cDNA with the RevertAid first strand cDNA synthesis kit (Thermo, USA, K1622) according to the manufacturer’s instructions. Real-time PCR was prepared with SYBR qPCR Premix (Takara, Japan, RR420L), and then performed with a real-time PCR detection instrument (Bio–Rad CFX96, USA). The primers used for real-time PCR were purchased from Tsingke Biotech (Beijing, China), and the sequences are listed in Table 1. The mRNA expression was normalized to GAPDH, and fold changes were calculated using the ΔΔCt method.

Immunoprecipitation and immunoblotting

Cultured cells were lysed in RIPA buffer supplemented with protease inhibitors (Roche, NJ, USA). Then, the cell lysates were subjected to SDS–PAGE separation and were transferred to PVDF membranes. The membranes were subjected to immunoblot assays using antibodies against CRB3 (Santa Cruz, USA, sc-292449, both isoforms CRB3a and CRB3b can be detected), GCP2 (Santa Cruz, USA, sc-377117), GCP3 (Santa Cruz, USA, sc-373758), GCP4 (Santa Cruz, USA, sc-271876), GCP5 (Santa Cruz, USA, sc-365837), GCP6 (Abcam, UK, ab95172), γ-tubulin (Proteintech, USA, #66320-1-Ig), Rab11 (Abcam, UK, ab128913), GFP (Roche, USA, #11814460001), Flag (Sigma, USA, F7425), CEP290 (Santa Cruz, USA, sc-390637), GSK3-β (CST, USA, #12456), β-catenin (CST, USA, #8480), GAPDH (Proteintech, USA, #HRP-6004), and β-actin (Proteintech, USA, #HRP-60008). HRP-conjugated secondary antibodies (CST, USA, #7074, #7076) were then incubated at room temperature for 1 h in the dark. Final detection was achieved by using ECL Plus (Millipore, Germany, WBULS0500). Proteins for the immunoprecipitation assay were lysed with NP40 buffer. The cell lysates were detected by a Dynabeads protein G immunoprecipitation kit (Invitrogen, USA, #10007D) following the manufacturer’s instructions. The elution protein complexes were subjected to an immunoblot assay.

Protein identification and bioinformatics analyses

LC-MS/MS analysis was conducted by PTM Bio (Zhejiang, China). Q ExactiveTM Plus (Thermo, MA, USA) was used for tandem mass spectrometry data analysis. The Blast2GO program against the UniProt database was used to perform functional annotation and classification of the identified proteins. Subsequently, DAVID tools were used to perform pathway enrichment analysis.

Scanning electron microscope (SEM) observations

MCF10A cells with CRB3-GFP and CRB3 knockdown were plated on culture slides, and doxycycline (Dox) was added to induce the expression of CRB3-GFP. After the cells were fully confluent, the slides were fixed with 4% paraformaldehyde at 4°C overnight, followed by 1% osmium tetroxide at 4°C for 1 h. Dehydration was performed at room temperature, and the samples were dried using the critical point drying method. Images were obtained by using a scanning electron microscope (Hitachi TM1000, Japan).

Clinical data

Patient data, breast cancer tissues, and adjacent para-cancerous tissues were collected from the First Affiliated Hospital of Xi’an Jiaotong University (Shaanxi, China). All patients signed informed consent forms before surgery. This research was authorized by the Ethics Committee of the First Affiliated Hospital of Xi’an Jiaotong University and conducted in conformity with the Declaration of Helsinki.

Statistics

Statistical analysis was performed using SPSS statistics 23.0 for Windows (IBM, Armonk, USA). All experiments were repeated at least three times. The values are expressed as the mean ± standard deviation (SD). Unpaired Student’s t test was used to compare the differences between two groups. One-way ANOVA followed by Dunnett’s multiple comparisons test was used for multiple comparisons. The χ2 test was used to assess the significance of the observed frequencies. P<0.05 was considered an indicator of statistical significance.

Supporting information